# Homework 1: Data Types, Functions and Conditionals (23 points)

name: Ningyuan Wang

email: wningyua@umich.edu

I did not discuss this homework with anyone.
This homework assignment took me 6 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. Thus, even though we may ask you to check for errors in your functions, you should not include in your submission any examples of your functions actually raising those errors.

<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not comprehensive, so you are encouraged to read the problem carefully and verify your code covers all possible cases.

## Problem 1: Severance textbook exercises (8 points) 
Complete the following exercises from the textbook which you can find [here](https://www.py4e.com/html3/) for free 
- Ch 1, Exercises 3,5,6,8  
- Ch 2, Exercises 3,4  
- Ch 3, Exercises 1, 2  
- Ch 4, Exercise 6  
- Ch 6, Exercise 1  

---
### Chapter 1, Exercises (2 points) 

#### Exercise 3 
What is the difference between a compiler and an interpreter?

<!-- grade_id: ch1-ex3-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# For example: "ans = 'This is a sample answer.'"
# You can use markdown formatting. 
# YOUR CODE HERE
ans = ("A compiler handles the entire program in a file, and then it runs a process to translate the high-level source code into machine language. "
       "In contrast, a interpreter (such as Python) reads the source code of the program directly, line by line. "
       "Compilers are generally faster than interpreters.")
# raise NotImplementedError()

```

<!-- grade_id: ch1-ex3-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

#### Exercise 5
What is wrong with the following code:
```
>>> primt 'Hello world!'
File "<stdin>", line 1
  primt 'Hello world!'
                     ^
SyntaxError: invalid syntax
>>>
```

<!-- grade_id: ch1-ex5-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = ("The code has violated the grammar rules of Python. Specifically, there is no parenthesis around the string. "
"Besides, the function 'print' has wrong spelling and it will raise a NameError later.")

# raise NotImplementedError()


```

<!-- grade_id: ch1-ex5-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

#### Exercise 6
Where in the computer is a variable such as `x` stored after the following Python line finishes?

a) Central processing unit

b) Main Memory

c) Secondary Memory

d) Input Devices

e) Output Devices

<!-- grade_id: cell-b0f556c01b7e2276, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'. 
# Should be one of ['a', 'b', 'c', 'd', 'e']
# YOUR CODE HERE
ans = 'd'
# raise NotImplementedError()
```

<!-- grade_id: cell-3176d2d916422842, points: 0.5, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

#### Exercise 8
Explain each of the following using an example of a human capability: (1) Central processing unit, (2) Main Memory, (3) Secondary Memory, (4) Input Device, and (5) Output Device. For example, “What is the human equivalent to a Central Processing Unit”?

<!-- grade_id: ch1-ex8-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = ("(1). CPU is the heart of any computer. It runs the software that we write."
        " (2). Main Memory stores programs and data. But it loses its information when the power is turned off. It's like short-term memory in human."
          " (3). Secondary Memory is also used to store information, but it is much slower than the main memory. The advantage of the secondary memory is that it can store information even when there is no power to the computer. Examples of secondary memory are disk drives or flash memory. It's kind of like long-memory in human. "
              " (4) and (5). Input and Output Devices are simply our screen, keyboard, mouse, microphone, speaker, touchpad, etc. They are all of the ways we interact with the computer. They'are like the way of human interacting with the world, such as books are an input device and reading reports are an output deivce for students.")
# raise NotImplementedError()
```

<!-- grade_id: ch1-ex8-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

### Chapter 2, Exercises (2 points) 

#### Exercise 3
 Write a function to prompt the user for hours and rate per hour to compute gross pay.

<!-- grade_id: cell-312caeaaf6abe1a9, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    return float(hours)*float(rate_per_hour)
    #raise NotImplementedError()
```

<!-- grade_id: cell-458cb3b5b6333645, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(3, 10) == 30
```

#### Exercise 4
Assume that we execute the following assignment statements:
```
width = 17
height = 12.0
```
For each of the following expressions, write the value of the expression and the type (of the value of the expression).

1. `width//2`

2. `width/2.0`

3. `height/3`

4. `1 + 2 * 5`

```python
'''width = 17
height = 12.0
type(width//2)
type(width/2.0)
type(height/3)
type(1 + 2 * 5)'''

```

Output:
```
'width = 17\nheight = 12.0\ntype(width//2)\ntype(width/2.0)\ntype(height/3)\ntype(1 + 2 * 5)'
```

<!-- grade_id: cell-e69aa106b6a89086, points: 0, grade: false, solution: true -->
```python
# Enter you answer as python list of tuples (value, type) in a variable 'ans'.
# For example, "ans = [(1,int), (2.0, float)]" (not actual answer)
# where the order of the list corresponds to order above.

# YOUR CODE HERE
ans = [(8,int), (8.5,float), (4.0,float), (11,int)]
#raise NotImplementedError()
```

<!-- grade_id: cell-796c8c2828e551ce, points: 1, grade: true, solution: false -->
```python
assert ans[0] == (8,int)
```

---
### Chapter 3, Exercises (2 points) 

#### Exercise 1 
Rewrite your `get_gross_pay` function from ch2, exercise 3 to give the employee 1.5 times the hourly rate for hours worked above 40 hours.

<!-- grade_id: cell-8f610a982f866fac, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    if hours <= 40:
        return hours*rate_per_hour
    else:
        return 40*rate_per_hour+(hours-40)*(1.5*rate_per_hour)
    #raise NotImplementedError()
```

<!-- grade_id: cell-537d9d1c5dcd290f, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(0, 1000)==0
assert get_gross_pay(2, 10)==20
```

#### Exercise 2
Rewrite your `get_gross_pay` program from ch2, exercise 3 using `try` and `except` so that your program handles non-numeric input gracefully by printing a message and exiting the program. The following shows two executions of the program:

```
>>> get_gross_pay(20, 'nine')
Error, please enter numeric input

>>> get_gross_pay('forty', 'nine')
Error, please enter numeric input
```

<!-- grade_id: cell-b57ce5b94cb6cf92, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    #raise NotImplementedError()
    try:
        return float(hours)*float(rate_per_hour)
    except:
        print("Error, please enter numeric input")
  
```

<!-- grade_id: cell-582c5cae10097fd1, points: 1, grade: true, solution: false -->
```python
try: 
    get_gross_pay('nine', 'five')
    get_gross_pay('nine', 5)
    get_gross_pay(5, 5)
    get_gross_pay(5, 'nine')
    assert True
except: 
    assert False
```

Output:
```
Error, please enter numeric input
Error, please enter numeric input
Error, please enter numeric input
```

### Chapter 4, Exercises (1 point)

#### Exercise 1 
What is the purpose of the `def` keyword in Python?

a) It is slang that means “the following code is really cool”

b) It indicates the start of a function

c) It indicates that the following indented section of code is to be stored for later

d) b and c are both true

e) None of the above

<!-- grade_id: ch4-ex1-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'. 
# Should be one of ['a', 'b', 'c', 'd', 'e']
# YOUR CODE HERE
#raise NotImplementedError()
ans = "b"
```

<!-- grade_id: ch4-ex1-test, points: 1, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

### Chapter 6, Exercises (1 point) 

#### Exercise 6
Read the documentation of the string methods at https://docs.python.org/library/stdtypes.html#string-methods You might want to experiment with some of them to make sure you understand how they work. `strip` and `replace` are particularly useful.

<!-- grade_id: cell-9a484821ec48c9fc, points: 0, grade: false, solution: true -->
```python
# Create a variable 'ans' with the string: 'I read the documentation.'
# YOUR CODE HERE
ans = "I read the documentation."
# raise NotImplementedError()
```

<!-- grade_id: cell-4dd3c9dad6369e20, points: 1, grade: true, solution: false -->
```python
assert ans == 'I read the documentation.'
```

## Problem 2: Euclid's Algorithm (3 points)

Euclid’s algorithm (https://en.wikipedia.org/wiki/Euclidean_algorithm) is a method
for finding the greatest common divisor (GCD) of two integers. Recall that the GCD of
two integers $m$ and $n$ is the largest number that divides both $m$ and $n$.

---
### Part 1 (2 points)

The __Wikipedia__ page above includes several pseudocode implementations of Euclid’s
algorithm. Choose one of these, and use it to implement a function `gcd`, which takes
two integers as its arguments and returns their GCD. You may assume that both
inputs are integers, so there is no need to include any error checking in your function.

**Note:** this is one of the rare occasions where you have my explicit permission to
look up your answer. Unless otherwise stated (e.g., as in this problem), looking up
solutions on Wikipedia or in any other non-class resource will be considered cheating!

<!-- grade_id: gcd, points: 0, grade: false, solution: true -->
```python
def gcd(a,b):
    """Compute the great common divisor (GCD) of two integers a,b. 
    """
    # YOUR CODE HERE
   
    while b != 0:
        t = b
        b = a%b
        a = t
    return a

# raise NotImplementedError()

```

<!-- grade_id: gcd_tests, points: 2, grade: true, solution: false -->
```python
"""Check that your gcd function returns the correct output for several inputs"""
assert gcd(4,5) == 1
assert gcd(20,10) == 10
```

### Part 2 (1 point) 
What does your function do if one or both of the arguments are negative?

<!-- grade_id: cell-9b7a530cbce57b8d, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = "It will return a negative number, but it doesn't make sense of the definition of gcd (i.e. gcd should be a postive number). To make the function more reasonable, I probably need to modify the previous function so that it can handle the negative numbers as well."
#raise NotImplementedError()
```

<!-- grade_id: cell-2bdb84d0d76cdb07, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

## Problem 3: Approximating Euler's number $e$ (8 points) 

The base of the natural logarithm, $e$, is typically deﬁned as the inﬁnite sum

$$
    e = \sum_{k=0}^{\infty} \frac{1}{k!} = 1 + 1 + \frac{1}{2} + \frac{1}{6} + \frac{1}{24}
$$

where $k!$ denotes the factorial of $k$, 

$$
k! = k \cdot (k-1) \cdot (k-2) ... 3 \cdot 2 \cdot 1 
$$

where we deﬁne $0! = 1$ by convention. For more on Euler’s number, see [here](https://en.wikipedia.org/wiki/E_(mathematical_constant)). In this problem, we will explore diﬀerent approaches to approximating this number.

### Part 1 (2 points) 

An early characterization of Euler’s number, due to Jacob Bernoulli, was as the limit of

$$ 
(1 + \frac{1}{x})^{x}
$$

as $x \rightarrow \infty$. Deﬁne a function called `euler_limit` that takes as an argument an integer $n$, and returns a `ﬂoat` that approximates $e$ by taking $x = n$ in Equation (2). You may assume that the input to your function will be a positive integer (of course, your function will still run just ﬁne if `n` is, say, a positive ﬂoat, but we will only use integer values in what follows).

<!-- grade_id: euler_limit, points: 0, grade: false, solution: true -->
```python
def euler_limit(n): 
    """Takes in a positive integer n and returns a float approximating the natural number e."""
    approx = (1+ 1/n)**n
    return approx
```

<!-- grade_id: test_euler_limit1, points: 1, grade: true, solution: false -->
```python
"""Check that your euler_limit function returns the correct output for several inputs"""
assert euler_limit(1) == 2.0 
assert euler_limit(2) == 2.25
```

<!-- grade_id: test_euler_limit2, points: 1, grade: true, solution: false -->
```python
assert round(euler_limit(3),3) == 2.370
```

### Part 2 (1 point) 
What happens when you call `euler_limit(n)` for really huge values of `n` (say, $10^{16}$ or $10^{18}$ )? Why does this happen? 

**Hint:** the answer has to do with ﬂoating point arithmetic. Think about how $1/n$ is represented on your computer when $n$ is big.

<!-- grade_id: prob3-part2-ans, points: 0, grade: false, solution: true -->
```python
print(euler_limit(10**16))
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = "It returns 1.0. Mathematically, when n is large enough, 1/n approaches to zero, and therefore the limit of e approaches to 1. In Python, there are 53 bits of precision available for a Python float. For numbers larger than that, Python keeps the number of digits manageable by displaying a rounded value instead. Therefore, Python keeps a rounded value automatically in the situations of really huge values. "
# raise NotImplementedError()
```

Output:
```
1.0
```

<!-- grade_id: prob3-part2-tests, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

### Part 3 (2 points) 
Deﬁne a function called `euler_sum` that takes a single non-negative integer argument `n`, and returns an approximation to $e$ based on the ﬁrst `n` terms of the sum in Equation (1). Your function should return a ﬂoat. 
You may assume that the input will be a non-negative integer, so you do not need to include error checking in your function. 

As an example, `euler_sum(4)` should return the sum of the ﬁrst four terms in Equation 1, so that `euler_sum(4)` returns $1 + 1 + 1/2 + 1/6 ≈ 2.667$. 

**Note:** the sum in Equation 1 starts counting with $k = 0$ (i.e., it is “0-indexed”), while our function starts counting with $n = 1$ (i.e., it is “1-indexed”). `euler_sum(1)` should use one term from Equation (1), so that `euler_sum(1)` returns `1`. Similarly, `euler_sum(0)` should return `0`, since by convention an empty sum is equal to zero. 

**Note:** you may use the math.factorial function to compute $k!$, but I recommend, for the sake of practice, implementing the factorial function yourself, as it is a nice example of a problem that is easily solved with recursion.

<!-- grade_id: euler-sum, points: 0, grade: false, solution: true -->
```python
def factorial(k): # suppose n is a non-negative number
    factorial = 1
    if k == 0:
        return factorial
    else:
        for i in range(1, k+1):
            factorial = factorial*i
        return factorial
              

def euler_sum(n): # n is a non-negative integer
    sum = 0 
    if n == 0:
        return sum
    else: # n>0
        for i in range(0,n):
            f_inverse = 1/factorial(i)
            sum = sum + f_inverse
        return sum


```

<!-- grade_id: test-euler-sum1, points: 1, grade: true, solution: false -->
```python
assert euler_sum(0) == 0 
assert round(euler_sum(4),3) == 2.667
```

<!-- grade_id: test-euler-sum2, points: 1, grade: true, solution: false -->
```python
assert euler_sum(3) == 2.5
```

### Part 4 (1 point) 
Deﬁne a function called `euler_approx` that takes a single argument, a positive ﬂoat `epsilon`, and uses the sum in (1) to obtain an approximation of $e$ that is within `epsilon` of the true value of $e$. 

**Hint:** use a while-loop. 

**Note:** you can use the Python math module to get the true value of $e$ (up to ﬂoating point accuracy): `math.exp(1)`.

<!-- grade_id: euler-approx, points: 0, grade: false, solution: true -->
```python
import math

def euler_approx(epsilon):
    epsilon>0 and isinstance(epsilon, float)# check the input: a postive float
    if True:
        n = 1 #start from the first term of euler sum
        while abs(euler_sum(n)- math.exp(1))>= epsilon:
            n = n+1
    return euler_sum(n)
```

<!-- grade_id: euler-approx-tests, points: 1, grade: true, solution: false -->
```python
assert abs(euler_approx(1e0) - math.exp(1)) < 1e0
```

### Part 5 (1 point) 
Deﬁne functions called `print_euler_sum_table` and `print_euler_lim_table` that each takes a single positive integer `n` as an argument and prints the error between the approximation given by, respectively, `euler_sum(k)` or `euler_limit(k)` and the true value of $e$ as `k` ranges from `1` to `n`, one per line. That is, the functions should print a table of the the approximation error for diﬀerent choices of `k`.

<!-- grade_id: prob3-ex5-func, points: 0, grade: false, solution: true -->
```python
def print_euler_sum_table(k):# a positive integer
    # YOUR CODE HERE
    for i in range(1, k+1):
        print(abs(euler_sum(i) - math.exp(1)))

def print_euler_lim_table(k):# a positive integer
    # YOUR CODE HERE
    for i in range(1, k+1):
        print(abs(euler_limit(i) - math.exp(1)))

```

<!-- grade_id: cell-fccdb9f21de162ca, points: 1, grade: true, solution: false -->
```python
# Run this cell to check your functions. 
print_euler_sum_table(10)
print()
print_euler_lim_table(10)
```

Output:
```
1.718281828459045
0.7182818284590451
0.2182818284590451
0.05161516179237857
0.009948495125712054
0.0016151617923787498
0.0002262729034896438
2.7860205076724043e-05
3.0586177750535626e-06
3.0288585284310443e-07

0.7182818284590451
0.4682818284590451
0.3479114580886753
0.2768755784590451
0.22996182845904567
0.19665545671693163
0.17178213141833298
0.1524973145086972
0.13710703674584668
0.12453936835904278
```

### Part 6 (1 point) 
Which approximation is better? 

<!-- grade_id: prob3-part6-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = "As n is greater 2, euler sum has a smaller error, thus a better approximation. Otherwise, euler limit is better."
# raise NotImplementedError()
```

<!-- grade_id: prob3-part6-test, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

## Problem 4: Testing properties of integers (5 points)
In this problem, you’ll get a bit more practice working with conditionals, and a ﬁrst exposure to the kind of thinking that is required in a typical “coding interview” question. A positive integer $n$ is a power of 2 if $n = 2^p$ for some integer $p \geq 0$.

### Part 1 (2 points) 
Write a function `is_power_of_2` that takes a positive integer as its only argument and returns a Boolean indicating whether or not the input is a power of 2. 

You may assume that the input is a positive integer. You may not use the built-in `math.sqrt` function in your solution. You should need only the division and modulus (%) operations. 

**Hint:** the simplest solution to this problem makes use of recursion, though recursion is not necessary.

<!-- grade_id: is_power_of_2, points: 0, grade: false, solution: true -->
```python
def is_power_of_2(n):# input is a postive integer
    if n == 1:
        return True
    elif n%2 == 1:
        return False
    else:
        return is_power_of_2(n/2)
```

<!-- grade_id: is_power_of_2_test1, points: 1, grade: true, solution: false -->
```python
assert is_power_of_2(1)
assert is_power_of_2(2)
assert not is_power_of_2(15)
```

<!-- grade_id: is_power_of_2_test2, points: 1, grade: true, solution: false -->
```python
assert is_power_of_2(4)
```

### Part 2 (3 points)
Generalize your previous solution to a function `is_power` that takes two positive integers as its arguments, `b` and `n`, in that order, and returns a Boolean. `is_power(b,n)` should return `True` if `n` is a power of `b` and `False` otherwise.

<!-- grade_id: is_power, points: 0, grade: false, solution: true -->
```python
def is_power(b,n):
    # YOUR CODE HERE
    if b ==1:
        if n ==1:
            return True
        else:
            return False
    else: 
        if n == 1:
            return True
        elif n%b in range(1, b):
            return False
        else:
            return is_power(b,n/b)
```

<!-- grade_id: is_power_test1, points: 1.5, grade: true, solution: false -->
```python
assert is_power(5, 125)
assert not is_power(5, 150)
```

<!-- grade_id: is_power_test2, points: 1.5, grade: true, solution: false -->
```python
assert not is_power(1, 2)
assert is_power(3, 1)
```

```python

```