# Homework 5 (35 points + 3 extra credits)

name: Ningyuan Wang

email: wningyua@umich.edu

This homework assignment took me 13 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).)

I discussed Extra 1 with Ming-Chen Lu; Extra 2 with Jie Cao. 

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. Before submitting, please make sure make sure __Cells->Run__ executes without errors; errors in your code cause problems for our auto-grader. Thus, even though we may ask you to check for errors in your functions, you should not include in your submission any examples of your functions actually raising those errors.

## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. Additional tests will be used at grading time, you are encouraged to read the problem carefully and verify your code covers all appropriate cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:

```python
%load_ext pycodestyle_magic
%flake8_on --ignore E266 --max_line_length 99
```

Output:
```
The pycodestyle_magic extension is already loaded. To reload it, use:
  %reload_ext pycodestyle_magic
```

```python
import numpy as np
import itertools
import functools
import random
import math
import scipy.special
```

### Problem 1: Basic Review Questions (10 points)

1. <b>(2 points)</b> Write a function called `collatz` takes a positive integer `num` as its argument, and returns a list. Your function should perform appropriate error checking for it. The list is generated as follows: The first element of the list is `num`. Continue doing the following things: if the last element of the list, `number`, is even, then the next element of the list is `number` / 2; if the last element of the list, `number`, is odd, then the next element of the list is 3 * number + 1. Stop when the last element of the list is 1. (Note: We guarantee the length of the list is finite.)

<!-- grade_id: cell-108419d75c8da43a, points: 0, grade: false, solution: true -->
```python
def collatz(num):  # positive integer , error checking
    if not isinstance(num, int):
        raise TypeError("The input should be an integer.")
    if num <= 0:
        raise ValueError("The input should be a postive integer.")
    ls = []
    ls.append(num)

    while num != 1:
        if num % 2 == 0:  # even
            num = num/2
            ls.append(num)
        else:  # odd
            num = 3*num+1
            ls.append(num)
    return(ls)
```

<!-- grade_id: collatz, points: 0.5, grade: true, solution: false -->
```python
assert collatz(3) == [3, 10, 5, 16, 8, 4, 2, 1]
```

<!-- grade_id: collatz-value-error, points: 0.5, grade: true, solution: false -->
```python
assert type(collatz(5)) == list
```

<!-- grade_id: collatz-type-error, points: 0.5, grade: true, solution: false -->
```python
try:
    collatz(-1)

except ValueError:
    pass

else:
    raise AssertionError("Function allows for negative int as input but it shouldn't.")
```

<!-- grade_id: cell-d6cb9403c8069e69, points: 0.5, grade: true, solution: false -->
```python
assert collatz(2) == [2, 1]
```

2. <b>(1.5 points)</b> __Casino Game:__ This game in the casino consists of tossing a die with N faces twice. There is a number written on each face of the die (these numbers are not necessarily distinct). In order to win, you must get the number $A$ on the first toss and the number $B$ on the second toss of the die. The excited viewers want to know the probability that you will win the game. Define a function `casino_game` that takes a list of integers  `n`, an integer `a` and an integer `b` as its arguments and returns the probability that you win the game. Your function should perform appropriate error checking. Assume that you get each face of the die with the same probability on each toss and that tosses are mutually independent.

<!-- grade_id: cell-611c0aafe5caa0cc, points: 0, grade: false, solution: true -->
```python
def casino_game(n, a, b):
    if not isinstance(n, list):
        raise TypeError("Please enter a number list.")
    if not isinstance(a, int):
        raise TypeError("Please enter an integer number")
    if not isinstance(b, int):
        raise TypeError("Please enter an integer number")
    if not all(isinstance(x, int) for x in n):
        raise TypeError("Please enter an integer list.")
    if a not in n or b not in n:
        return 0
    else:  # a and b are both in n
        p1 = n.count(a)/len(n)
        p2 = n.count(b)/len(n)
        return p1*p2
```

<!-- grade_id: casino-game-easy, points: 0.5, grade: true, solution: false -->
```python
assert casino_game([1, 2, 3, 4], 1, 1) == 0.0625
```

<!-- grade_id: cell-42f199155837eb9e, points: 0.5, grade: true, solution: false -->
```python
assert casino_game([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 0, 0) == 1.0
```

<!-- grade_id: casino-game-type-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    casino_game(0, 1, 1)

except TypeError:
    pass

else:
    raise AssertionError("Function allows for int input where a list is expected.")
```

3. <b>(2 points)</b> Write a function `dict_comp` that takes a positive integer `n` as its argument and returns a dictionary that contains the numbers (between 1 and `n`) as its keys and the first `n` composite numbers as its values. Your function should perform appropriate error checking for it. Recall that a composite number is any integer $p > 1$ which is not prime; therefore, 4 is the first composite number.

<!-- grade_id: cell-5129b62a1c02655b, points: 0, grade: false, solution: true -->
```python
def composites():
    (comp, num) = ([], 1)
    while True:
        (comp, num) = (comp, num+1)
        if any(num % x == 0 for x in range(2, num)):
            comp.append(num)
            yield num


def dict_comp(n):  # postive integer
    if not isinstance(n, int):
        raise TypeError("n is not an integer.")
    if n <= 0:
        raise ValueError("n is not a positive integer.")
    dic = {}
    ans = composites()
    composite_seq = [next(ans) for _ in range(n)]
    for i in range(1, n+1):
        dic[i] = composite_seq[i-1]
    return dic
```

<!-- grade_id: dict-comp-easy, points: 0.5, grade: true, solution: false -->
```python
assert dict_comp(5) == {1: 4, 2: 6, 3: 8, 4: 9, 5: 10}
```

<!-- grade_id: dict-comp-hard, points: 0.5, grade: true, solution: false -->
```python
assert type(dict_comp(3)) == dict
```

<!-- grade_id: dict-type-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    dict_comp([0])

except TypeError:
    pass

else:
    raise AssertionError("Function allows for string input where an int is expected.")
```

<!-- grade_id: dict-value-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    dict_comp(0)

except ValueError:
    pass

else:
    raise AssertionError("Value allows for input that is not a positive integer.")
```

4. <b>(1 point)</b> Write a one-line expression that computes the sum of the first 10 even square numbers (starting from $2^2$). For ease of grading, please assign the output of this expression to a variable called `even_squares`. Please make sure that the cell that defines even_squares is the one directly above the test case for this question.

<!-- grade_id: even-squares, points: 0, grade: false, solution: true -->
```python
even_squares = sum([n**2 for n in range(1, 11)])
```

<!-- grade_id: cell-6f5cd818d00af6fd, points: 1, grade: true, solution: false -->
```python
# Make sure that the cell that defines even_squares is directly above this one.
assert type(even_squares) == int or type(even_squares) == float
```

5. <b>(1 point)</b> Write a generator expression that enumerates the tetrahedral numbers. The n-th tetrahedral number ($n = 1,\ 2,\ \cdots$) is given by $Tn = \binom{n+2}{3}$， where $\binom{x}{y}$ is the binomial coefficient
$$\binom{x}{y} = \frac{x!}{y!(x-y)!}$$
For ease of grading, please assign this generator expression to a variable called `tetra`. Hint: you may find it useful to define a generator that enumerates the positive integers.

<!-- grade_id: cell-2ade7aa604144668, points: 0, grade: false, solution: true -->
```python
import math


def pos_int():
    n = 0
    while True:
        n += 1
        yield n


x = pos_int()
tetra = (math.factorial(n + 2) / (math.factorial(3) * math.factorial(n - 1))for n in x)

# [next(tetra) for _ in range(10)]
```

<!-- grade_id: tetra, points: 1, grade: true, solution: false -->
```python
import types
assert isinstance(tetra, types.GeneratorType)
```

6. <b>(1 point)</b> Write a one-line expression that computes the geometric mean of the first 12 tetrahedral numbers. You may use the generator that you wrote in the previous problem. Re-calling that generator does not count as contributing to the one line requirement. Recall that the geometric mean of a collection of n numbers $a_1,\ a_2,\ \cdots,\ a_n$ is given by ($\Pi^n_{i=1} a_i)^{\frac{1}{n}}$. For ease of grading, please assign the output of this expression to a variable called `tetra_geom`. Please leave the cell that defines tetra_geom directly above the cell that contains the test case for this problem. 

<!-- grade_id: tetra-geom, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
tetra_geom = functools.reduce(lambda x, y: x*y, [next(tetra) for _ in range(12)])**(1/12)  
```

<!-- grade_id: tetra-geom-tests, points: 1, grade: true, solution: false -->
```python
# Make sure that tetra_geom is created in the cell directly above this one.
assert type(float(tetra_geom)) == float
```

7. <b>(1.5 points)</b> You are given a string `s`, which contains alphanumeric characters only (this means you can assume that the string will have no spaces). Your task is to sort the string `s` in the following manner:  
 - All sorted lowercase letters are ahead of uppercase letters.
 - All sorted uppercase letters are ahead of digits.
 - All sorted odd digits are ahead of sorted even digits.

Write a function called `sort_str`, which takes the string `s` as its argument, and return the result. Your function should perform appropriate error checking for the input.

<!-- grade_id: cell-752e6972babe6b23, points: 0, grade: false, solution: true -->
```python
def sort_str(s):
    if not isinstance(s, str):
        raise TypeError("The input is not a string.")
    low = "".join(sorted([c for c in s if c.islower()]))
    upp = "".join(sorted([c for c in s if c.isupper()]))
    dig = "".join(sorted([c for c in s if c.isdigit()]))
    odd = "".join(sorted([d for d in dig if int(d) % 2 != 0]))
    even = "".join(sorted([d for d in dig if int(d) % 2 == 0]))
    return low + upp + odd + even
```

<!-- grade_id: sort-str, points: 1, grade: true, solution: false -->
```python
assert sort_str('Sorting1234') == 'ginortS1324'
```

<!-- grade_id: sort-str-type-error, points: 0.5, grade: true, solution: false -->
```python
try:
    sort_str(123,)

except TypeError:
    pass

else:
    raise AssertionError("Function allows for int input where string inputs is expected.")
```

### Problem 2: Fantasy Game (5 points)
1. <b>(2 points)</b> __Fantasy Game Inventory:__ You are creating a fantasy video game. The data structure to model the player’s inventory will be a dictionary where the keys are string values describing the item in the inventory and the value is an integer value detailing how many of that item the player has. For example, the dictionary value `{'rope': 1, 'torch': 6, 'gold coin': 42, 'dagger': 1, 'arrow': 12}` means the player has 1 rope, 6 torches, 42 gold coins, and so on. Write a function named `displayInventory` that takes a dictionary `inventory` as its argument and prints out the inventory line by line. Using the example dictionary, it should print it out EXACTLY like the following:

Inventory: <br>
12 arrow <br>
42 gold coin <br>
1 rope <br>
6 torch <br>
1 dagger <br>
Total number of items: 62 <br>

__Note:__ Your function should raise TypeError if the input is not a dictionary.

__Note:__ You will not be graded on the order of the inventory items. The only print order that matters is that "Inventory:" is printed before all of the items and the total is printed after all of the items. Everything in between can occur in any order so long as it is accurate.

<!-- grade_id: cell-d6b22abac7bcf9ea, points: 0, grade: false, solution: true -->
```python
def display_inventory(inventory):  # error check
    if not isinstance(inventory, dict):
        raise TypeError("Please check your inventory.")
    print("Inventory:\n")
    total = 0
    for k, v in inventory.items():
        total = total + v
        print("%d %s \n" % (v, k))
    print("Total number of items: %d" % total)
```

<!-- grade_id: display-inventory, points: 1, grade: true, solution: false -->
```python
inventory = {'rope': 1, 'torch': 6, 'gold coin': 42, 'dagger': 1, 'arrow': 12, 'map fragments': 3}
assert display_inventory(inventory) is None
```

Output:
```
Inventory:

1 rope 

6 torch 

42 gold coin 

1 dagger 

12 arrow 

3 map fragments 

Total number of items: 65
```

<!-- grade_id: display-error, points: 1, grade: true, solution: false -->
```python
try:
    display_inventory(["1: rope", "2: torch"])

except TypeError:
    pass

else:
    raise AssertionError("Input of list doesn't raise proper error.")
```

2. <b>(3 points)</b> __List to Dictionary Function for Fantasy Game Inventory:__ Imagine that a vanquished dragon’s loot is represented as a list of strings like this: `dragonLoot = ['gold coin', 'dagger', 'gold coin', 'gold coin', 'ruby']`. Write a function named `addToInventory(inventory, added_items)`, where the `inventory` argument is a dictionary representing the player’s inventory (like in the previous question) and the `added_items` argument is a list like dragonLoot. Your function should check to make sure that the types of the input are appropriate (e.g., that inventory is a dictionary and addedItems should be a list of strings). The `addToInventory` function should return a dictionary that represents the updated inventory. __Note:__ the addedItems list can contain multiples of the same item.

<!-- grade_id: cell-d0f5b2d73004c4e5, points: 0, grade: false, solution: true -->
```python
def add_to_inventory(inventory, added_items):
    if not isinstance(inventory, dict):
        raise TypeError("Please check your inventory.")
    if not isinstance(added_items, list):
        raise TypeError("Please check your added items.")
    if not all(isinstance(x, str) for x in added_items):
        raise TypeError("Please check your items.")

    for item in added_items:
        if item in inventory:  # in the dict keys
            inventory[item] = inventory[item] + 1
        else:
            inventory[item] = 1
    return inventory
```

<!-- grade_id: add-to-inv-easy, points: 1, grade: true, solution: false -->
```python
inv = {'gold coin': 42, 'rope': 1}
dragon_loot = ['gold coin', 'dagger', 'ruby']
assert add_to_inventory(inv, dragon_loot) == {'dagger': 1, 'gold coin': 43, 'rope': 1, 'ruby': 1}
```

<!-- grade_id: add-to-inv-hard, points: 1, grade: true, solution: false -->
```python
inv = {'gold coin': 42, 'rope': 1}
dragon_loot = ['gold coin', 'dagger', 'gold coin', 'gold coin', 'ruby']
correct_dict = {'dagger': 1, 'gold coin': 45, 'rope': 1, 'ruby': 1}
assert add_to_inventory(inv, dragon_loot) == correct_dict
```

<!-- grade_id: add-to-inv-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    add_to_inventory([1, 1], ['a', 'b'])

except TypeError:
    pass

else:
    raise AssertionError("First input is not a dict.")
```

<!-- grade_id: add-to-inv-errors2, points: 0.5, grade: true, solution: false -->
```python
try:
    add_to_inventory({'a': 1, 'b': 1}, [2, 3])

except ValueError:
    pass
except TypeError:
    pass

else:
    raise AssertionError("Second input is a list of ints where a list of strings is expected.")
```

### Problem 3: More Fun with Strings (5 points)
1. <b>(2 points)</b> __Merge the Tools:__ Consider the following: A string, $s$, of length $n$  where $s=c_0c_1 \ldots c_{n-1}$. An integer, $k$ , where $k$ is a factor of $n$. We can split $s$ into $n/k$ subsegments where each subsegment, $t_i$ , consists of a contiguous block of $k$ characters in $s$. Then, use each $t_i$ to create string $u_i$ such that:  
 - The characters in $u_i$ are a subsequence of the characters in $t_i$.  
 - Any repeat occurrence of a character is removed from the string such that each character in $u_i$ occurs exactly once. In other words, if the character at some index $j$  in $t_i$  occurs at a previous index $< j$ in $t_i$, then do not include the character in string $u_i$.
  
 Write a function called `str_seg`, which takes the string `s` and the value `k` as its two argument, and return a list containing $\frac{n}{k}$ elements, where $i$-th element denotes the string $u_i$. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking. This means you can assume s and k will be of correct types in all the test cases and k will be a positive integer.

<!-- grade_id: cell-ad502ca49fb8f47e, points: 0, grade: false, solution: true -->
```python
def str_seg(s, k):
    n = len(s)
    lst = []
    for i in range(0, n, k):
        t = ''.join((s[i:i+k]))
        if k != 1:
            u = t[0]

            for i in range(1, len(t)):
                if t[i] not in u:
                    u = u + t[i]
                    lst.append(u)
        else:
            lst = lst + [t]
    return lst
```

<!-- grade_id: empty-string, points: 0.5, grade: true, solution: false -->
```python
assert type(str_seg("This is the first test case", 16)) == list
```

<!-- grade_id: str-seg, points: 1, grade: true, solution: false -->
```python
assert str_seg('AABCAAADA', 3) == ['AB', 'CA', 'AD']
```

<!-- grade_id: str-seg-edge-cases, points: 0.5, grade: true, solution: false -->
```python
assert str_seg("01234567890", 1) == ['0', '1', '2', '3', '4', '5', '6', '7', '8', '9', '0']
```

2. <b>(3 points)</b> __Simple Arithmetics:__ You are given multiple sequences of button presses of a simple pocket calculator that consist of digits and arithmetic operators. For each such sequence find the number it would produce on a pocket calculator's display. __Note that the pocket calculator evaluates the operators in the order in which they are given (i.e., there is no operator precedence (PEMDAS doesn't apply here)).__ Assume that the display of the calculator is large enough to show the result, and that its memory is sufficient to store all intermediate results. Write a function called `simple_arith`, which takes a string as its only argument, and returns the result as a float. In this problem, A string represents one sequence of button presses for a pocket calculator. The sequence consists of non-negative integers and arithmetic operators and ends with an equal sign. __It can also contain spaces to improve readability.__ __The operators only contain "+-*/"__, and the operator / represents integer division, rounded down. You may assume that no test case contains division by zero. __Hint:__ you may use `eval()` function. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking. This means you can assume that the input will always be a string with a valid mathematical expression in it.

<!-- grade_id: cell-c83de32d388931f8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def simple_arith(s):
    s = s.replace(" ", "")  # remove space
    s = s.split('=')[0]  # get the string before =
    num = []
    operator_ind = [i for i in range(len(s)) if not s[i].isdigit()]
    operator_ind.append(0)
    operator_ind.append(len(s))
    operator_ind = sorted(operator_ind)
    for i in range(len(operator_ind)-1):
        n = s[operator_ind[i]:operator_ind[i+1]]
        num.append(n)
#     print(num)
    t = num[0]
    for i in range(1, len(num)):
        t = math.floor(eval(t + num[i]))
        t = str(t)
    return float(t)  
```

<!-- grade_id: correct-return-type, points: 0.5, grade: true, solution: false -->
```python
assert type(simple_arith("5*5=")) == float
```

<!-- grade_id: simple-arith-simple, points: 0.5, grade: true, solution: false -->
```python
assert simple_arith('1+2+3 + 3 + 1*2 =') == 20.0
### BEDIN HIDDEN TESTS
assert simple_arith('1 / 2 + 2 - 1 * 3 =') == 3.0
assert simple_arith('2 - 1 + 1 / 2 * 3 =') == 3.0
```

<!-- grade_id: simple-arith-digits, points: 0.5, grade: true, solution: false -->
```python
assert simple_arith('100-99-10+5=') == -4.0
```

<!-- grade_id: equals-only, points: 0.5, grade: true, solution: false -->
```python
assert simple_arith('9-9/5-5=') == -5.0
```

<!-- grade_id: simple-arith-hard, points: 1, grade: true, solution: false -->
```python
assert simple_arith("5*5/4 =") == 6.0
```

### Problem 4: More Practice (15 points)

1. <b>(2 points)</b> You are given a string $s$. Suppose a character '$c$' occurs consecutively $x$ times in the string. __In this problem, we assume each character is in '0123456789', representing a number between 0 and 9.__ Replace these consecutive occurrences of the character '$c$' with $(x,c)$ in the string and store into a list. Write a function called `comp_str`, which takes the string `s` as its argument, and return the list containing the replacements. For example, the result of `comp_str('1222311')` is `[(1, 1), (3, 2), (1, 3), (2, 1)]`. Reason: First, the character '1' occurs only once, so it is replaced by (1,1). Then the character '2' occurs three times, and it is replaced by (3,2) and so on. __Hint:__ _itertools.groupby_ might be helpful in this question. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking.

<!-- grade_id: cell-988506a92b718ef2, points: 0, grade: false, solution: true -->
```python
def comp_str(s):
    groups = []
    keys = []
    ls = [c for c in s]
    for k, g in itertools.groupby(ls):
        keys.append(int(k))
        groups.append(len(list(g)))
    return [(groups[i], keys[i]) for i in range(len(keys))]
```

<!-- grade_id: comp-str-return, points: 0.5, grade: true, solution: false -->
```python
assert type(comp_str('989898989')) == list
```

<!-- grade_id: comp-str-easy, points: 0.5, grade: true, solution: false -->
```python
assert comp_str('1222311') == [(1, 1), (3, 2), (1, 3), (2, 1)]
```

<!-- grade_id: comp-str-hard, points: 0.5, grade: true, solution: false -->
```python
assert comp_str('2') == [(1, 2)]
```

<!-- grade_id: comp-str-empty, points: 0.5, grade: true, solution: false -->
```python
assert comp_str('030303') == [(1, 0), (1, 3), (1, 0), (1, 3), (1, 0), (1, 3)]
```

2. <b>(2 points)</b> __The Mirror of Galadriel:__ With water from the stream Galadriel filled the basin to the brim, and breathed on it, and when the water was still again she spoke. 'Here is the Mirror of Galadriel,' she said. 'I have brought you here so that you may look in it, if you will. For this is what your folk would call magic, I believe; though I do not understand clearly what they mean; and they seem also to use the same word of the deceits of the Enemy. But this, if you will, is the magic of Galadriel. Did you not say that you wished to see Elf-magic?' - Galadriel to Frodo and Sam, describing her Mirror. 

 We call a string S magical if every substring of S appears in Galadriel's Mirror (under lateral inversion). In other words, a magical string is a string where every substring has its reverse in the string. Write a function `magical`, that takes a string `s` as its argument, determine if it is magical or not, and return the result: "Yes" is it is magical, or "No" if it is not magical. For example, `magical('aba')` is 'Yes'. Reason: the list of substrings are : `a, b, ab, ba, aba`. The reverse of each of these strings is present as a substring of S too. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking.

<!-- grade_id: cell-513c33ab888a566f, points: 0, grade: false, solution: true -->
```python
def magical(s):
    # YOUR CODE HERE
    subs = []
    revs = []
    for i in range(len(s)):
        for j in range(i+1, len(s)+1):
            sub = s[i: j]
            subs.append(sub)
#     print(subs)
    for s in subs:
        rev = s[::-1]
        revs.append(rev)
#     print(revs)
    if all(rev in subs for rev in revs):
        return "Yes"
    else:
        return "No"

```

<!-- grade_id: magical-empty, points: 0.5, grade: true, solution: false -->
```python
assert type(magical("random string")) == str
```

<!-- grade_id: magical-standard, points: 1, grade: true, solution: false -->
```python
assert magical('aba') == 'Yes'
```

<!-- grade_id: magical-hard, points: 0.5, grade: true, solution: false -->
```python
assert magical('magical') == 'No'
```

3. <b>(3 points)</b> Everybody hates Raymond. He’s the largest stamp collector on planet earth and because of that he always makes fun of all the others at the stamp collector parties. Fortunately everybody loves Lucy, and she has a plan. She secretly asks her friends whether they could lend her some stamps, so that she can embarrass Raymond by showing an even larger collection than his. Raymond is so sure about his superiority that he always tells how many stamps he’ll show.And since Lucy knows how many she owns, she knows how many more she needs. She also knows how many friends would lend her some stamps and how many each would lend. But she’s like to borrow from as few friends as possible and if she needs too many then she’d rather not do it at all. Can you tell her the minimum number of friends she needs to borrow from?

 Write a function called `stamps`, which contains the following arguments: how many stamps Lucy needs to borrow and a list containing the number of stamps each of her friends is offering. This function should return the minimum number of friends Lucy needs to borrow stamps from as an integer. If it’s impossible even if she borrows everything from everybody, return "Impossible!" instead. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking.

<!-- grade_id: cell-84a1624b777aea7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def stamps(num, lst):
    offer = sorted(lst, reverse=True)
    s = 0
    count = 0
    for i in range(len(offer)):
        s = s + offer[i]
        count = count + 1
        if s >= num:
            return count
    if s < num:
        return "Impossible!"
```

<!-- grade_id: stampls-long, points: 1, grade: true, solution: false -->
```python
assert type(stamps(4, [1, 1, 1])) == str
assert type(stamps(4, [5])) == int
```

<!-- grade_id: stample-basic, points: 1, grade: true, solution: false -->
```python
assert stamps(100, [13, 17, 42, 9, 23, 57]) == 3
```

<!-- grade_id: stamps-ones, points: 1, grade: true, solution: false -->
```python
assert stamps(4, [1, 1, 1, 1, 1]) == 4
```

4. <b>(2 points)</b> You are given a grid with R rows (numbered 1 through R) and C columns (numbered 1 through C). A cell in row r and column c is denoted by (r, c). Two cells in the grid are adjacent if they have a common side. For each valid i and j, there is a value $a_{i, j}$ written in cell (i, j). A cell in the grid is stable if the number of cells in the grid which are adjacent to this cell is strictly greater than the value written in this cell. The whole grid is stable if all cells in the grid are stable.

    Define a function `is_stable` that takes an r-by-c n numpy array (i.e., an ndarray) as its only argument and returns a Python Boolean that is True if and only if the argument is a stable array.

<!-- grade_id: cell-4e8532c0e87d3f45, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_stable(a):
    if not isinstance(a, np.ndarray):
        return False
    if len(np.shape(a)) != 2:
        return False
    (r, c) = a.shape
    adj = np.ones((r, c))*8  # assume each entry has 8 adjancent cells
    for i in range(r):
        for j in range(c):
            if (i == 0 and j == 0) or (i == r-1 and j == 0) or \
               (i == 0 and j == c-1) or (i == r-1 and j == c-1):
                adj[i, j] = 3
            elif i == 0 or i == r - 1:
                adj[i, j] = adj[i, j] - 3
            elif j == 0 or j == c - 1:
                adj[i, j] = adj[i, j] - 3
    comparison = adj > a
    return comparison.all()
```

<!-- grade_id: basic-example, points: 1, grade: true, solution: false -->
```python
array1 = np.array([1, 0, 0,
                   0, 0, 0,
                   0, 0, 0]).reshape(3, 3)
array2 = np.array([3, 0, 0,
                   0, 0, 0,
                   0, 0, 0]).reshape(3, 3)


assert is_stable(array1)
assert not is_stable(array2)
```

<!-- grade_id: 3-by-3, points: 0.5, grade: true, solution: false -->
```python
assert not is_stable(np.array([1., 0., 1.]))
```

<!-- grade_id: 4-by-4, points: 0.5, grade: true, solution: false -->
```python
assert is_stable(np.array([[1, 2, 2, 1], [2, 3, 3, 1], [2, 3, 3, 2], [1, 2, 2, 1]]))
```

5. <b>(3 points)</b> The Fibonacci sequence $F_0,\ F_1,\cdots$ is a special infinite sequence of non-negative integers, where $F_0=0,\ F_1=1$ and for each integer $n \ge 2$, $F_n=F_{n−1}+F_{n−2}$. Consider the sequence $D$ of the last decimal digits of the first $N$ Fibonacci numbers, i.e. $D=(F_{0\%10},F_{1\%10},…,F_{(N−1)\%10})$. Now, you should perform the following process: 
 - Let $D=(D_1,D_2,\cdots,D_l)$.
 - If $l=1$, the process ends.
 - Create a new sequence $E=(D_2,D_4,…,D_{2⌊l/2⌋})$. In other words, $E$ is the sequence created by removing all odd-indexed elements from $D$. 
 - Change $D$ to $E$. 
 
 When this process terminates, the sequence $D$ contains only one number. You have to find this number. Define a function `find_D` that takes a positive integer as its only argument and returns the only number in $D$. Your function should perform appropriate error checking. __Hints:__ Fibonacci's final digits cycle every 60 numbers.

<!-- grade_id: cell-e19142dafbe0c5bd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class D_generator():
    '''Fibo sequence'''
    def __init__(self):
        self.f0 = 0
        self.f1 = 1

    def __next__(self):
        self.n = self.f0
        (self.f0, self.f1) = (self.f1, self.f0+self.f1)
        return (self.n % 10)

    def __iter__(self):
        return(self)


def find_D(n):
    if not isinstance(n, int):
        raise TypeError("Please enter an integer.")
    if n <= 0:
        raise ValueError("Please enter a positive integer.")
    ans1 = D_generator()
    d = [next(ans1) for _ in range(n)]
    e = ["i", 'n']
    while len(e) > 1:
        ls = len(d)
        if ls == 1:
            return 0  # If  𝑙=1 , the process ends.
        else:
            e = [v for i, v in enumerate(d) if i % 2 != 0]
            d = e
    return e[0]

```

<!-- grade_id: D-edge, points: 0.5, grade: true, solution: false -->
```python
assert type(find_D(4)) == int
```

<!-- grade_id: D-basic, points: 1, grade: true, solution: false -->
```python
assert find_D(11) == 3
assert find_D(5) == 2
```

<!-- grade_id: D-difficult, points: 0.5, grade: true, solution: false -->
```python
assert find_D(71) == 2
```

<!-- grade_id: cell-353f2157b9618d06, points: 0.5, grade: true, solution: false -->
```python
try:
    find_D(1.5)

except TypeError:
    pass

else:
    raise AssertionError("Function doesn't raise TypeError for float input")
```

<!-- grade_id: D-negative, points: 0.5, grade: true, solution: false -->
```python
try:
    find_D(0)

except ValueError:
    pass

else:
    raise AssertionError("Function doesn't raise ValueError for input of 0")
```

6. <b>(3 points)</b> You are given a function $f(X) = X^2$. You are also given $K$ lists. The $i$-th list consists of $N_i$ elements. You have to pick one element from each list so that the value from the equation below is maximized:
$$S=\left(f\left(X_{1}\right)+f\left(X_{2}\right)+\ldots+f\left(X_{k}\right)\right) \% M$$
 $X_i$ denotes the element picked from the $i$-th list. Find the maximized value of $S$ obtained. \% denotes the modulo operator.   Note that you need to take exactly one element from each list, not necessarily the largest element. You add the squares of the chosen elements and perform the modulo operation. The maximum value that you can obtain, will be the answer to the problem.

 Write a function called `max_mod`, which takes, `M` and $K$ lists as its arguments, and return the maximized value of $S$.

 For example, `max_mod(1000, [2,5,4], [3,7,8,9], [5,7,8,9,10])` is `206`.  Reason: Picking 5 from the 1-st list, 9 from the 2-nd list and 10 from the 3-rd list gives the maximum S value equal to $\left(5^{2}+9^{2}+10^{2}\right) \% 1000 = 206$. __Hint:__ `itertools.product` might be helpful.

<!-- grade_id: cell-33f9473d081b4ac1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def max_mod(M, *K):
    k_lst = [k for k in K]
    combs = [e for e in itertools.product(*k_lst)]
    s = []
    for i in range(len(combs)):
        res = functools.reduce(lambda x, y: x+y, map(lambda x: x**2, combs[i]))
        res = res % M
        s.append(res)
    return max(s)
```

<!-- grade_id: maxmod-basic, points: 1, grade: true, solution: false -->
```python
assert max_mod(10, [2, 6, 3]) == 9
```

<!-- grade_id: maxmod-edge, points: 1, grade: true, solution: false -->
```python
assert max_mod(1, [2]) == 0
```

<!-- grade_id: maxmod-hard, points: 1, grade: true, solution: false -->
```python
assert max_mod(2, [3, 4], [4, 5], [5, 7]) == 1
```

### Problem 5: Extra Credit (3 points)

1. <b>(+1 EC point)</b> __Coin Flip Streaks:__ Suppose you flip a coin 100 times and write down an “H” for each heads and “T” for each tails, you’ll create a list that looks like “T T T T H H H H T T.” If you ask a human to make up 100 random coin flips, you’ll probably end up with alternating headtail results like “H T H T H H T H T T,” which looks random (to humans), but isn’t mathematically random. A human will almost never write down a streak of several heads or several tails in a row, even though it is highly likely to happen in truly random coin flips. Humans are predictably bad at being random.

 Write a function called `coin_streaks`, which takes a positive integer `n` as its only argument. Your function should perform appropriate error checking. Your function breaks up the experiment into two parts: the first part generates a list of randomly selected 'heads' and 'tails' values, and the second part checks if there is a streak in it. Put all of this code in a loop that repeats the experiment 10,000 times so we can find out what percentage of the coin flips contains a streak of $n$ heads or $n$ tails in a row. Return this percentage. As a hint, the function call `random.randint(0, 1)` will return a 0 value 50% of the time and a 1 value the other 50% of the time.

<!-- grade_id: cell-a082c7044b6be7a3, points: 0, grade: false, solution: true -->
```python
def coin_streaks(n):
    if not isinstance(n, int):
        raise TypeError("Please enter an integer.")
    if n <= 0:
        raise ValueError("Please enter a positive integer.")

    streak = 0
    r = 0

    while r < 10000:  # repeat 10000 times
        lst = []

        for i in range(100):
            if random.randint(0, 1) == 0:
                lst.append("H")
            else:
                lst.append("T")

        count = 0
        for i in range(len(lst)-1):
            if lst[i] == lst[i+1]:
                count += 1
            else:
                count = 0
            if count == n-1:
                break
        if count == n-1:
            streak = streak + 1
        r = r + 1
    return (streak / 10000) * 100
```

<!-- grade_id: coin-streaks, points: 1, grade: true, solution: false -->
```python
# If this cell passes on our end you will get 1 extra credit point.
```

2. <b>(+1 EC point)</b> The only printer in the computer science students' union is experiencing an extremely heavy workload. Sometimes there are a hundred jobs in the printer queue and you may have to wait for hours to get a single page of output.

 Because some jobs are more important than others, the Hacker General has invented and implemented a simple priority system for the print job queue. Now, each job is assigned a priority between 1 and 9 (with 9 being the highest priority, and 1 being the lowest), The first position in the queue is number 0, the second is number 1, and so on. The printer operates as follows.

 - The first job J in queue is taken from the queue.
 - If there is some job in the queue with a higher priority than job J, then move J to the end of the queue without printing it.
 - Otherwise, print job J (and do not put it back in the queue).

 In this way, all those important muffin recipes that the Hacker General is printing get printed very quickly. Of course, those annoying term papers that others are printing may have to wait for quite some time to get printed, but that's life.

 Your problem with the new policy is that it has become quite tricky to determine when your print job will actually be completed. You decide to write a program to figure this out. The program will be given the current queue (as a list of priorities) as well as the position of your job in the queue, and must then calculate how long it will take until your job is printed, assuming that no additional jobs will be added to the queue. To simplify matters, we assume that printing a job always takes exactly one minute, and that adding and removing jobs from the queue is instantaneous.

 Write a function called `printer_queue`, which contains the following arguments: the position of your job, and a list containing integers in the range 1 to 9 giving the priorities of the jobs in the queue (The first integer gives the priority of the first job, the second integer the priority of the second job, and so on). This function should return a number , which is the number of minutes until your job is completely printed, assuming that no additional print jobs will arrive.

<!-- grade_id: cell-f05d8e72a8786e8b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def printer_queue(pos, queue):
    if not isinstance(pos, int):
        raise TypeError
    if not isinstance(queue, list):
        raise TypeError
    if not all(isinstance(q, int) for q in queue):
        raise TypeError
    if not all(q > 0 and q < 10 for q in queue):
        raise ValueError
    job = queue[pos]
    before = 0
    for i in range(len(queue)):
        if queue[i] == job and i < pos:
            before = before + 1
    pri_lst = sorted(queue, reverse=True)
    return pri_lst.index(job) + 1 + before


# assert printer_queue(0, [2, 7, 9, 2, 5, 6]) == 5
# assert printer_queue(1, [2, 7, 9, 2, 5, 6]) == 2
# assert printer_queue(0, [8, 7, 2, 5, 3]) == 1
# assert printer_queue(2, [8, 7, 2, 5, 3]) == 5
# assert printer_queue(1, [2, 2, 2, 2, 2]) == 2
```

<!-- grade_id: printer-queue, points: 1, grade: true, solution: false -->
```python
# If this cell passes on our end you will get 1 extra credit point.
```

3. <b>(+1 EC point)</b> Johnny has some difficulty memorizing the small prime numbers. So, his computer science teacher has asked him to play with the following puzzle game frequently. The puzzle is a 3x3 board consisting of numbers from 1 to 9. The objective of the puzzle is to swap the tiles until the following final state is reached:

 1 2 3

 4 5 6

 7 8 9

 At each step, Johnny may swap two adjacent tiles if their sum is a prime number. Two tiles are considered adjacent if they have a common edge. Help Johnny to figure out whether the given 3-by-3 board can reach the goal state or not. Define a function `solve_puzzle` that takes an 3-by-3 numpy array (i.e., an ndarray) as its only argument and returns Python Boolean that is True if and only if the given 3-by-3 board can reach the goal state. __Hint:__ You can start from the goal state, walk through all the possible states and save them into a global variable. You function just need to judge if the given state is inside the possible states or not.

<!-- grade_id: cell-26806f821d6284c0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: solve-puzzle, points: 1, grade: true, solution: false -->
```python
# If this cell passes on our end you will get 1 extra credit point.
```